```python
# Initialize Otter
import otter
grader = otter.Notebook("q1.ipynb")
```

# Qiskit Assignment 1
## Practice with Pauli Gates and other Qiskit tools
Welcome to your second Qiskit assignment!

### Learning Objectives
1. Build Pauli gates from Qiskit's U gate
2. Use the U gate to reverse a series of operations
3. Work with various rotations on the bloch sphere
4. Understand IBM's circuit transpilation process

**NAME**: REPLACE WITH YOUR NAME

**STUDENT ID**: REPLACE WITH YOUR STUDENT ID

```python
# Import Qiskit and other needed packages
from qiskit import *
from qiskit.visualization import plot_histogram
from qiskit.quantum_info import Statevector
from qiskit.circuit.random import random_circuit
import numpy as np
import pprint
import random
```

#### Task 1 - Constructing Pauli Z gate from u Gate (10 pts)
We can use [Qiskit's U Gate](https://qiskit.org/documentation/stubs/qiskit.circuit.QuantumCircuit.u.html#qiskit.circuit.QuantumCircuit.u) to construct arbitrary quantum operations. Fill in the function below to return `qc_pauli_z`, a QuantumCircuit satisfying the following conditions:
- it has 1 qubit, [initialized](https://qiskit.org/documentation/stubs/qiskit.circuit.QuantumCircuit.initialize.html#qiskit.circuit.QuantumCircuit.initialize) to the parameter `initial_state`
- it has 1 classical bit
- it has a U gate with parameters which perform the same rotation as a Pauli Z gate
- it does not use the built in Z gate
- it performs a measurement following the rotation

```python
def qc_pauli_z(initial_state=[1,0]):
    ...
```

```python
qc_pauli_z().draw(output='mpl')
```

```python
grader.check("Task 1")
```

#### Task 2 (2A, 2B, 2C) - Unitary Inverse Puzzles (30 pts)

We'll study the idea of [uncomputation](https://qiskit.org/textbook/ch-algorithms/grover.html#5.2-Uncomputing,-and-Completing-the-Oracle) during our discussion of quantum algorithms. In general, we may find it helpful to return a qubit to its initial state. 

This process is typically straightforward due to the properties of unitary gates. However, your task is to do so using only a single U gate. Complete the partial circuits below such that the measurements will yield a state equivalent to `initial_state` *up to a global phase*.

#### Task 2A

```python
def reverse_a(initial_state=[1,0]):
    qc = QuantumCircuit(1,1)
    qc.initialize(initial_state)
    qc.x(0)
    qc.h(0)
    qc.y(0)
    qc.x(0)
    qc.z(0)
    qc.barrier()
    
   ...

    qc.barrier()
    qc.measure(0,0)
    return qc
```

```python
reverse_a().draw(output='mpl')
```

```python
grader.check("Task 2A")
```

#### Task 2B

```python
def reverse_b(initial_state=[1,0]):
    qc = QuantumCircuit(1,1)
    qc.initialize(initial_state)
    for i in range(5):
        qc.x(0)
        qc.y(0)
        qc.z(0)
        qc.h(0)
    qc.barrier()
    
    
    ...

    qc.barrier()
    qc.measure(0,0)
    return qc
```

```python
reverse_b().draw(output='mpl')
```

```python
grader.check("Task 2B")
```

#### Task 2C

**Hint**: The P gate generalizes rotation about the Z-axis to an arbitrary angle $\phi$, where
$P(\phi)=\begin{pmatrix} 1 & 0 \\ 0 & e^{i\phi} \end{pmatrix}$

```python
def reverse_c(initial_state=[1,0]):
    qc = QuantumCircuit(1,1)
    qc.initialize(initial_state)
    qc.x(0)
    for i in range(1,6):
        qc.p((-1)**(i)*np.pi/(2**i), 0)
    qc.z(0)
    qc.barrier()
    
    ...

    qc.barrier()
    qc.measure(0,0)
    return qc
```

```python
reverse_c().draw(output='mpl')
```

```python
grader.check("Task 2C")
```

<!-- BEGIN QUESTION -->

#### Task 3 (3A, 3B, 3C) - Using Rotation to Obtain Probabilities (30 pts)

#### Task 3A 
Fill in the function below to return `qc_rot_a`, a single-qubit QuantumCircuit satisfying the following conditions:
- it performs a measurement to a single classical bit
- Pr(seeing $|0\rangle$ on measurement) = `0.75`
- your circuit only uses gates from the following list: X, Y, Z, P, H, U

Plot your results using a histogram to verify your solution over `1024` trials.

```python
def qc_rot_a():
    ...
```

```python
qc_rot_a().draw(output='mpl')
```

```python
# Plot your results in this cell!

...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Task 3B - Rotation Operator Gates
Again, fill in the function below to return `qc_rot_b`, a single-qubit QuantumCircuit satisfying the following conditions:
- it performs a measurement to a single classical bit
- Pr(seeing $|0\rangle$ on measurement) = `0.75`
- your circuit only uses gates from the following list: [RX, RY, RZ](https://en.wikipedia.org/wiki/Quantum_logic_gate#Rotation_operator_gates)

Plot your results using a histogram to verify your solution over `1024` trials.

```python
def qc_rot_b():
    ...
```

```python
qc_rot_b().draw(output='mpl')
```

```python
# Plot your results in this cell!

...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Task 3C
Suppose we apply a Z gate to your circuit from task 3B just before measuring. How will the probability of measuring $|0\rangle$ change from that of the original circuit? Will measurement on the modified circuit yield a state equivalent to the original circuit up to a global phase?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Task 4 - P vs. RZ (10 pts)
In tasks 2C, the P gate was introduced. In lecture and task 3B, you've seen the RZ gate. Both are related to Z-axis rotations, but what utility is there to having each? Is there a mathematical relationship between them? If so, describe it. 

Feel free to use any resources for your research, including `lecture 6: A single qubit`.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Task 5 (5A, 5B) - Transpiling Circuits (20 pts)
When you submit a job to IBM, the quantum computer will most likely run a different circuit than you built. This is because the quantum computer can only do a very limited set of operations relative to the number of unitary gates. For IBM devices, the transpile step reduces all single-qubit operations to I, X, [SX](https://qiskit.org/documentation/stubs/qiskit.circuit.library.SXGate.html), and RZ ([source](https://qiskit.org/textbook/ch-states/single-qubit-gates.html#7.-The-U-gate--)).

#### Task 5A
Choose one of the IBM backends (see the first assignment for a refresher on this). Use the [transpile](https://qiskit.org/documentation/stubs/qiskit.compiler.transpile.html) method to optimize the given circuit for the backend. Draw the transpiled circuit.

```python
qc_rand = random_circuit(1, 10, measure=True, seed=6)
qc_rand.draw(output='mpl')
```

```python
IBMQ.load_account()
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Task 5B
Which gates from {I, X, SX, RZ}, and how many of each, are used in the transpiled circuit?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Conclusion

Next week: the EV bomb!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for Task 1 (points: 10)
```
>>> qc_pauli_z().num_clbits == 1 and qc_pauli_z().num_qubits == 1
True
```
failure_message: Incorrect number of qubits or classical bits.
```
>>> def testNoQiskitZGate():
...     ops = qc_pauli_z().count_ops()
...     return 'z' not in ops
...     
>>> def testInitialState():
...     try:
...         num_initializations = qc_pauli_z().count_ops()['initialize']
...     except KeyError:
...         return False
...     else:
...         return num_initializations == 1
...     
>>> def testMeasurementPerformed():
...     try:
...         num_measurements = qc_pauli_z().count_ops()['measure']
...     except KeyError:
...         return False
...     else:
...         return num_measurements == 1
...     
>>> testMeasurementPerformed() and testNoQiskitZGate() and testInitialState()
True
```
failure_message: Circuit requirements not satisfied, see bullet list.
```
>>> def testRandomInitialState():
...     # let's choose something interesting as our initial state!
...     c = random.uniform(0, 2*np.pi)
...     initial_state = [(np.cos(c)), (np.sin(c))]
...     pauli_z_matrix = np.array(
...                                 [[1, 0],
...                                  [0,-1]]
...                              )
...     
...     qc = qc_pauli_z(initial_state)
...     qc.remove_final_measurements()
...     result_sv = Statevector.from_instruction(qc)
...     
...     return result_sv.equiv(Statevector(pauli_z_matrix.dot(initial_state)))
>>> 
>>> testRandomInitialState()
True
```
failure_message: Incorrect state returned.


Tests for Task 2A (points: 10)
```
>>> def testAcceptableGateUsage2A():
...     ops = reverse_a().count_ops()
...     x = ops['x'] == 2
...     y = ops['y'] == 1
...     z = ops['z'] == 1
...     h = ops['h'] == 1
...     try:
...         u = ops['u'] == 1
...     except KeyError:
...         u = False
...         
...     return all([x,y,z,h,u])
...     
>>> testAcceptableGateUsage2A()
True
```
failure_message: Circuit requirements not satisfied, check bullet list.
```
>>> def testKetZero2A():
...     qc = reverse_a()
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector([1,0]))
>>> 
>>> def testKetOne2A():
...     initial_state = [0,1]
...     qc = reverse_a(initial_state)
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector([0,1]))
...     
>>> testKetZero2A() and testKetOne2A()
True
```
failure_message: Basis states not passing.
```
>>> def testRandomSimple2A():
...     c = random.uniform(0, 2*np.pi)
...     initial_state = [(np.cos(c)), (np.sin(c))]
...     qc = reverse_a(initial_state)
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector(initial_state))
...     
>>> testRandomSimple2A()
True
```
failure_message: Random state not passing.


Tests for Task 2B (points: 10)
```
>>> def testAcceptableGateUsage2B():
...     ops = reverse_b().count_ops()
...     x = ops['x'] == 5
...     y = ops['y'] == 5
...     z = ops['z'] == 5
...     h = ops['h'] == 5
...     try:
...         u = ops['u'] == 1
...     except KeyError:
...         u = False
...     return all([x,y,z,h,u])
...     
>>> testAcceptableGateUsage2B()
True
```
failure_message: Circuit requirements not satisfied, check bullet list.
```
>>> def testKetZero2B():
...     qc = reverse_b()
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector([1,0]))
>>> 
>>> def testKetOne2B():
...     initial_state = [0,1]
...     qc = reverse_b(initial_state)
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector([0,1]))
...     
>>> testKetZero2B() and testKetOne2B()
True
```
failure_message: Basis states not passing.
```
>>> def testRandomSimple2B():
...     c = random.uniform(0, 2*np.pi)
...     initial_state = [(np.cos(c)), (np.sin(c))]
...     qc = reverse_b(initial_state)
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector(initial_state))
...     
>>> testRandomSimple2B()
True
```
failure_message: Random state not passing.


Tests for Task 2C (points: 10)
```
>>> def testAcceptableGateUsage2C():
...     ops = reverse_c().count_ops()
...     x = ops['x'] == 1
...     y = not 'y' in ops
...     z = ops['z'] == 1
...     h = not 'h' in ops
...     p = ops['p'] == 5
...     try:
...         u = ops['u'] == 1
...     except KeyError:
...         u = False
...     return all([x,y,z,h,p,u])
...     
>>> testAcceptableGateUsage2C()
True
```
failure_message: Circuit requirements not satisfied, check bullet list.
```
>>> def testKetZero2C():
...     qc = reverse_c()
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector([1,0]))
>>> 
>>> def testKetOne2C():
...     initial_state = [0,1]
...     qc = reverse_c(initial_state)
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector([0,1]))
...     
>>> testKetZero2C() and testKetOne2C()
True
```
failure_message: Basis states not passing.
```
>>> def testRandomSimple2C():
...     c = random.uniform(0, 2*np.pi)
...     initial_state = [(np.cos(c)), (np.sin(c))]
...     qc = reverse_c(initial_state)
...     qc.remove_final_measurements()
...     sv = Statevector.from_instruction(qc)
...     return sv.equiv(Statevector(initial_state))
...     
>>> testRandomSimple2C()
True
```
failure_message: Random state not passing.
