# ASSIGNMENT CONFIG
config_file: ../config.yml

## Qiskit Assignment 000
## README
Welcome to your first Quantum Computing assignment! This semester we are using Jupyter Notebooks to distribute assignments and collect your solutions. An advantage of this approach is that much of your work can be automatically graded.

```python
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
from IPython.display import display
```

```python
def gen(choices):
  return interactive (lambda ans:ans, ans=choices)
```

# BEGIN QUESTION
name: q1

What is $1+1$?

# BEGIN SOLUTION

```python
ans1 = gen([0,1,2])
display(ans1)
```

Output:
```
interactive(children=(Dropdown(description='ans', options=(0, 1, 2), value=0), Output()), _dom_classes=('widge…
```

```python
result1 = ans1.result
```

Output:
```
71
```

# END SOLUTION

# BEGIN TESTS

```python
assert result1 == 2
```

# END TESTS

# END QUESTION

Tests for Task 1 (points: 20)
```
>>> isinstance(simpleCircuit(), QuantumCircuit)
True
```
failure_message: EXPECTED FUNCTION TO RETURN TYPE QuantumCircuit
```
>>> simpleCircuit().num_clbits == 1
True
```
failure_message: EXPECTED CIRCUIT TO HAVE 1 CLASSICAL BIT
```
>>> simpleCircuit().num_qubits == 1
True
```
failure_message: EXPECTED CIRCUIT TO HAVE 1 QUBIT
```
>>> def testMeasurementPerformedSimple():
...     ops = simpleCircuit().count_ops()
...     return (ops['measure'] == 1 if 'measure' in ops else False)
>>> 
>>> testMeasurementPerformedSimple()
True
```
failure_message: EXPECTED CIRCUIT TO PERFORM A MEASUREMENT
```
>>> def testProbabilitiesSimple():
...     qc = simpleCircuit()
...     job = execute(qc, BasicAer.get_backend('statevector_simulator'), shots=1)
...     return list(job.result().get_statevector(qc)) == [1,0]
>>> 
>>> testProbabilitiesSimple()
True
```
failure_message: EXPECTED CIRCUIT TO PRODUCE STATE |0>
```
>>> def testMeasurementPerformedSimple():
...     ops = simpleCircuit().count_ops()
...     return (ops['measure'] == 1 if 'measure' in ops else False)
>>> def testProbabilitiesSimple():
...     qc = simpleCircuit()
...     job = execute(qc, BasicAer.get_backend('statevector_simulator'), shots=1)
...     return list(job.result().get_statevector(qc)) == [1,0]
>>> 
>>> isinstance(simpleCircuit(), QuantumCircuit) \
... and simpleCircuit().num_clbits == 1 \
... and simpleCircuit().num_qubits == 1 \
... and testMeasurementPerformedSimple() \
... and testProbabilitiesSimple()
True
```
failure_message: ONE OF THE ABOVE TESTS IS NOT PASSING


Tests for Task 3 (points: 10)
```
>>> '0' in counts and counts['0'] == 468
True
```
failure_message: EXPECTED simpleCircuit TO BE RUN 468 TIMES


Tests for Task 5 (points: 15)
```
>>> isinstance(opposites(), QuantumCircuit)
True
```
failure_message: EXPECTED FUNCTION TO RETURN TYPE QuantumCircuit
```
>>> opposites().num_clbits == 2
True
```
failure_message: EXPECTED CIRCUIT TO HAVE 2 CLASSICAL BITS
```
>>> opposites().num_qubits == 2
True
```
failure_message: EXPECTED CIRCUIT TO HAVE 2 QUBITS
```
>>> def testNoInitialize():
...     ops = opposites().count_ops()
...     return 'initialize' not in ops
>>> 
>>> testNoInitialize()
True
```
failure_message: CIRCUIT SHOULD NOT USE INITIALIZE
```
>>> def testRequirementsOpposites():
...     ops = opposites().count_ops()
...     if 'measure' not in ops:
...         return False
...     x = 'x' in ops and ops['x'] == 1
...     return ops['measure'] == 2 and x
>>> 
>>> testRequirementsOpposites()
True
```
failure_message: CIRCUIT SHOULD MEASURE EACH QUBIT TO ITS CORRESPONDING CLASSICAL BIT
```
>>> def testProbabilitiesOpposites():
...     qc = opposites().reverse_bits()
...     job = execute(qc, BasicAer.get_backend('statevector_simulator'), shots=10)
...     return list(job.result().get_statevector(qc)) == [0, 1, 0, 0]
>>> 
>>> testProbabilitiesOpposites()
True
```
failure_message: Measured unexpected state.
```
>>> 
>>> def testNoInitialize():
...     ops = opposites().count_ops()
...     return 'initialize' not in ops
>>> def testProbabilitiesOpposites():
...     qc = opposites().reverse_bits()
...     job = execute(qc, BasicAer.get_backend('statevector_simulator'), shots=10)
...     return list(job.result().get_statevector(qc)) == [0, 1, 0, 0]
>>> def testRequirementsOpposites():
...     ops = opposites().count_ops()
...     if 'measure' not in ops:
...         return False
...     x = 'x' in ops and ops['x'] == 1
...     return ops['measure'] == 2 and x
>>> 
>>> isinstance(opposites(), QuantumCircuit) \
... and opposites().num_clbits == 2 \
... and opposites().num_qubits == 2 \
... and testNoInitialize() \
... and testRequirementsOpposites() \
... and testProbabilitiesOpposites()
True
```
failure_message: ONE OF THE ABOVE TESTS IS NOT PASSING


Tests for Task 6 (points: 15)
```
>>> '01' in counts_oppo and counts_oppo['01'] == 468
True
```
failure_message: EXPECTED opposites TO BE RUN 468 TIMES
